# Lecture 4 Live Exercises

In Lecture 4 we covered
- Default Values for Functions
- Docstrings
- Abstraction
- Recursion

However, to make sure that our Python Greenhorns are getting enough practice, I've arranged problems slightly differently.

## Greenhorn Review Exercises

**1. Counting Characters: (2pts)** Write a function `char_count(x,y)` that takes in a string and a character and counts how many times that character shows up in that string. For example, for the string `s = 'njksdk#fj5j%jjabd#j'` and character `c = '#'`, `char_count(s,c)` should return `2`.

<!-- grade_id: cell-6da9b74021c6c2a6, points: 0, grade: false, solution: true -->
```python
def char_count(x,y):
    result = 0
    for i in x:
        if i == y:
            result = result +1
            #result += 1
    return result
```

```python
char_count('apple','e')
```

Output:
```
1
```

<!-- grade_id: cell-a4cdbbe8707ab8db, points: 2, grade: true, solution: false -->
```python
assert char_count('njksdk#fj5j%jjabd#j','#') == 2
assert char_count('apple','p') == 2
```

**2. Repeat Characters: (2 pts)** Write a function `repeat_chars(s,n)` that takes in a string `s` and positive integer `n` and repeats every character in that string *n* times. 

So `repeat_chars('Hello',4)` should return `'HHHHeeeelllloooo'` and
`repeat_chars('apple',2)` should return `aappppllee`.

<!-- grade_id: cell-fe60408b3548fccb, points: 0, grade: false, solution: true -->
```python
def repeat_chars(s,n):
    result = ''
    for i in s:
        result = result + n*i
    return result
```

```python
repeat_chars('gato',3)
```

Output:
```
'gggaaatttooo'
```

<!-- grade_id: cell-cca1acd1905f1a20, points: 2, grade: true, solution: false -->
```python
assert repeat_chars('Hello',4) == 'HHHHeeeelllllllloooo'
assert repeat_chars('BATMAN',3) == 'BBBAAATTTMMMAAANNN'
```

## Lecture 4 Exercises

**1. Print numbers, recursively: (2pts)** Write a function `countDown(n)` that uses recursion to return a string that counts down from `n`, separated by `...`, ending with the base case where `n==0`, which should return `Blast Off!`

In other words, `countDown(5)` should return the string `'5...4...3...2...1...Blast Off!'`.

<!-- grade_id: cell-faafbbf32ca6d2a3, points: 0, grade: false, solution: true -->
```python
def countDown(n):
    if n==0:
        return 'Blast Off!'
    else:
        return str(n) + "..." + countDown(n-1)
```

```python
countDown(3)
```

Output:
```
'3...2...1...Blast Off!'
```

<!-- grade_id: cell-a5ef999770143b85, points: 2, grade: true, solution: false -->
```python
assert countDown(3) == '3...2...1...Blast Off!'
assert countDown(10) == '10...9...8...7...6...5...4...3...2...1...Blast Off!'
```

**2. Triangular Numbers (two parts!): (3 pts)** 

**Part 1 (2pts)** Write a function called `triang(n)` that uses recursion to return the sum $$n+(n-1)+(n-2)+\cdots+ 1.$$

<!-- grade_id: cell-7a06125d6becd5a9, points: 0, grade: false, solution: true -->
```python
def triang(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-22e9d5c528c36298, points: 2, grade: true, solution: false -->
```python
assert triang(3) == 6
assert triang(10) == 55
```

**Q2: Part 2 (1pt- manually graded)** Add a docstring to the code of `triang(n)` so that when you call `help(triang)` it explains what the function does.

<!-- grade_id: cell-3e5db8cb418e85cc, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**3. Printing a string backwards and other tricks: (2pts)** 

For this problem, assume that you only know the following and a little bit of recursion.

- You can concatenate two strings using the `+` operation.
- If you want to know the length of a string `s`, you can use the command `len(s)`.
- If you want to access the $i^{th}$ entry in a string, you can use `s[i]`. Note that this starts counting at 0.
- The command `s[i:j]` returns the substring starting at entry number `i` and ends at `j-1`.

Use recursion to write a function `printBackwards(x)` that takes in a string and prints it backwards.

<!-- grade_id: cell-0bcda8a24408afd5, points: 0, grade: false, solution: true -->
```python
def printBackwards(x):
    if len(x) ==1:
        return x
    else:
        return x[len(x)-1] + printBackwards(x[:len(x)-1])
```

```python
printBackwards('cat')
```

Output:
```
'tac'
```

<!-- grade_id: cell-df0a7f89662a5b0a, points: 2, grade: true, solution: false -->
```python
assert printBackwards('taco') == 'ocat'
assert printBackwards('Was it a cat i saW') == 'Was i tac a ti saW'
```

**4. Testing whether an integer is prime relative to another: (4 pts)**

Write a function `isPrime` that takes in an integer `n` and returns whether it is prime or not by recursively checking divisibility by integers between 2 and $n/2$.

**HINT:** Use two arguments, where the second argument is the number you're checking divisibility by. Give this one a default value as in the code below.

<pre>
def isPrime(n,i=2):
    if i> n/2:
        return True
    # Your Solution here
</pre>

**Part 1 (2 pts):** Complete the function `isPrime` below.

<!-- grade_id: cell-62936e83ac51a4ce, points: 0, grade: false, solution: true -->
```python
def isPrime(n,i=2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e506b70638044aba, points: 2, grade: true, solution: false -->
```python
assert isPrime(3) == True
assert isPrime(17) == True
assert isPrime(20,6) == False
```

**Part 2 (2 pts):** Complete the following test function `testPrime(n)` that illustrates the result of calling `isPrime(i)` on `i=1,...,n`

<!-- grade_id: cell-549daf913cf15622, points: 2, grade: true, solution: true -->
```python
def testisPrime(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python

```