# Problem Set 03 (24 points)
## Due before the End of the Semester for Participation Credit
### Topics Covered
- Recursion
- Function Evaluation
- Operations on Strings
- Functions that return functions
- OOPing your Dream House

### **Exercise 1: Part A (2 pts, autograded):**

Recall that the Fibonacci sequence 0,1,1,2,3,5,8..., is defined by $f_0=0$, $f_1=f_2=1$ and $f_k = f_{k-1} + f_{k-2}$ for $k\geq 3$. 

In this problem we want to produce a **generalized Fibonacci sequence** where $g_0=0$, $g_1=n$, $g_2=m$ and $g_k = g_{k-1} + g_{k-2}$ for $k\geq 3$. 

For this problem, finish the function below `gen_Fib(n,m,k)` with a *recursive* implementation so that `gen_Fib(n,m,k)` returns the kth term in the generalized Fibonacci sequence, defined above.

<!-- grade_id: cell-a8d542d2301c9d18, points: 0, grade: false, solution: true -->
```python
def gen_Fib(n,m,k):
    """Takes in 3 integers n, m, k and 
    returns the kth element of the generalized Fibonacci sequence 
    that starts as ...0,n,m,n+m, n+2m,"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ef25a90574fa4c32, points: 2, grade: true, solution: false -->
```python
assert gen_Fib(1,1,7) == 13
assert gen_Fib(1,1,10) == 55
assert gen_Fib(1,3,5) == 11
assert gen_Fib(1,3,10) == 123
assert gen_Fib(1,-1,8) == -5
```

### **Exercise 1: Part B (2 pts, manually graded):**

For $n=2$ and $m=1$, the generalized Fibonacci sequence is called the **Lucas sequence**, which we will write $L_k$. 

Write some code below to calculate the limit $$\lim_{k\to\infty} \frac{L_{k+1}}{L_k} = ?$$

<!-- grade_id: cell-e3ad92a5a07389ea, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### **Exercise 1: Part C (4 pts, manually graded):**

How does your answer from Part B compare to the limit of $F_{k+1}(n,m)/F_k(n,m)$, where $F_k(n,m)$ denotes the generalized Fibonacci sequence for general seed values $n$ and $m$? Try at least 3 different pairs of values $(n,m)$, which are unique to you. Describe any patterns you see with Python comments to your code.

<!-- grade_id: cell-62a85f793ff2f805, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### **Exercise 2 (2 pts, autograded):** 
Finish the function below that computes "n double factorial", written $n!!$ , recursively. As a reminder, this is defined iteratively as $$n!!=n*(n-2)*(n-4)*...$$ Make sure your program handles even and odd cases. Assume $0!!=1$ and $1!!=1$.

<!-- grade_id: cell-e6e7a7efa7c7acef, points: 0, grade: false, solution: true -->
```python
def nDoubleFactorial(n):
    """Takes in an integer n and returns n!! 
    We assume n!! for n<=1 is 1."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e844acd308fab1a5, points: 2, grade: true, solution: false -->
```python
assert nDoubleFactorial(2) == 2
assert nDoubleFactorial(3) == 3
assert nDoubleFactorial(4) == 8
assert nDoubleFactorial(5) == 15
assert nDoubleFactorial(6) == 48
```

### **Exercise 3: (3 pts, manually graded): Counting recursive calls**

In the recursive function `fib` that we introduced in Lecture 04 to compute the Fibonacci numbers, how many times is `fib(2)` called when `fib(5)` is called? 

Write some code and explain your reasoning.

<!-- grade_id: cell-866d623b8a827e8c, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### **Exercise 4 (3pts, autograded): Functions that return Functions**

Write a function `myPoly(coeffs)` that takes in a list of coefficients `coeffs=[a_0, a_1, a_2,...,a_n]` and returns a *function*
$$f(x)= a_n x^n + \cdots + a_1 x + a_0$$

<!-- grade_id: cell-746de5091e43695e, points: 0, grade: false, solution: true -->
```python
def myPoly(coeffs):
    """coeffs is a list of floats or ints, function returns a polynomial with 
    the kth entry in coeffs as the coefficient of the monomial x^k"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ee2ed9943ecece8b, points: 3, grade: true, solution: false -->
```python
assert abs(myPoly([1,2,1])(5) - 36) <= .1
assert abs(myPoly([1,1,1])(5) - 31) <= .1
assert abs(myPoly([1,2,3])(4) - 57) <= .1
```

### **Exercise 6 (8 pts across two parts, manually graded)**

#### **Part A (2 pts, manually graded)**
Let's make a new class called `House` that where you have to provide the following attributes in the `__init__` function:
- address,
- style, (see [this list](https://www.diynetwork.com/how-to/rooms-and-spaces/exterior/26-popular-architectural-home-styles-pictures) for inspiration)
- yearBuilt
- numUnits (are there rental units in the property?)
- numFloors (how many floors does it have?)

You should also include the attributes and provide methods to change the values of (or just check!) 
- doorsLocked
- windowsClosed

<!-- grade_id: cell-03c4d0d2afb797d2, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### **Part B (6pts, manually graded)**

Add at least 3 new attributes and 3 methods for manipulating these attributes. Get creative! 

**Do your best to make sure these methods and attributes are unique to your implementation of the House class.**

<!-- grade_id: cell-c68aa375edd0c69b, points: 6, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```