# Problem Set 03 (20 points)
## Due Sunday September 26th, 2021 at 11:59PM
### Topics
- Recursion
- Function Evaluation
- Operations on Strings
- Functions that return functions
- OOPing your Dream House

### Exercise 1 (2 pts, autograded):

Recall that the Fibonacci sequence 0,1,1,2,3,5,8..., is defined by $f_0=0$, $f_1=f_2=1$ and $f_k = f_{k-1} + f_{k-2}$ for $k\geq 3$. 

In this problem we want to produce a **generalized Fibonacci sequence** where $g_0=0$, $g_1=n$, $g_2=m$ and $g_k = g_{k-1} + g_{k-2}$ for $k\geq 3$. 

For this problem, finish the function below `gen_Fib(n,m,k)` with a *recursive* implementation so that `gen_Fib(n,m,k)` returns the kth term in the generalized Fibonacci sequence, defined above.

<!-- grade_id: cell-2687064ceb2ed9da, points: 0, grade: false, solution: true -->
```python
def gen_Fib(n,m,k):
    """Takes in 3 integers n, m, k and 
    returns the kth element of the generalized Fibonacci sequence 
    that starts as ...0,n,m,n+m, n+2m,"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7d2f52820b706f74, points: 2, grade: true, solution: false -->
```python
assert gen_Fib(1,1,7) == 13
assert gen_Fib(1,1,10) == 55
assert gen_Fib(1,3,5) == 11
assert gen_Fib(1,3,10) == 123
assert gen_Fib(1,-1,8) == -5
```

### **Exercise 2 (2 pts, autograded):** 
Finish the function below that computes "n double factorial", written $n!!$ , recursively. As a reminder, this is defined iteratively as $$n!!=n*(n-2)*(n-4)*...$$ Make sure your program handles even and odd cases. Assume $0!!=1$ and $1!!=1$.

<!-- grade_id: cell-d3fc932800eb5b21, points: 0, grade: false, solution: true -->
```python
def nDoubleFactorial(n):
    """Takes in an integer n and returns n!! 
    We assume n!! for n<=1 is 1."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-3332b37b53ac2d5a, points: 2, grade: true, solution: false -->
```python
assert nDoubleFactorial(2) == 2
assert nDoubleFactorial(3) == 3
assert nDoubleFactorial(4) == 8
assert nDoubleFactorial(5) == 15
assert nDoubleFactorial(6) == 48
```

### Exercise 3: (3 pts, manually graded): Counting recursive calls

In the recursive function `fib` that we introduced in Lecture 04 to compute the Fibonacci numbers, how many times is `fib(2)` called when `fib(5)` is called? 

Write some code and explain your reasoning.

<!-- grade_id: cell-181b16ec2ec265a0, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Exercise 4: (5 pts, over two parts) 

#### Exercise 4, Part A (3 pts, autograded) 
Complete the `isPal` function below that takes in a string and returns `True` if the entered string is a palindrome and `False` otherwise. Punctuation marks, blank spaces, and capitalization are ignored, by virtue of the sub-function `toChar(s)`. Recall that a palindrome is a string that is the same whether it is read forwards or backwards. Your code should return true if it is given the string 'Was it a cat, I saw?'

<!-- grade_id: cell-7b0500bfaf000063, points: 0, grade: false, solution: true -->
```python
def isPalindrome(s): 
    """Assumes s is a str
    Returns True if the letters in s form a palindrome;
    False otherwise. Non-letters and capitalization are ignored."""
    def toChars(s): 
        t = s.lower() 
        letters = '' 
        for c in t:
            if c in 'abcdefghijklmnopqrstuvwxyz': 
                letters = letters + c
        return letters
    
    def isPal(s):
        # YOUR CODE HERE
        raise NotImplementedError()
    return isPal(toChars(s))
```

<!-- grade_id: cell-11891cc9059d5d91, points: 3, grade: true, solution: false -->
```python
assert isPalindrome('Was it a cat, I saw?') == True
assert isPalindrome("Never odd or even") == True
```

#### Exercise 4, Part B (2pts, manually graded) 
Explain in words why the code for `toChar(s)` had to create a new string `t` with `t= s.lower()`. Why couldn't we modify the string `s` directly?

<!-- grade_id: cell-a2afc668e2da74c6, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

### Exercise 5 (3pts, autograded): Functions that return Functions

Write a function `myPoly(coeffs)` that takes in a list of coefficients `coeffs=[a_0, a_1, a_2,...,a_n]` and returns a *function*
$$f(x)= a_n x^n + \cdots + a_1 x + a_0$$

<!-- grade_id: cell-cbe8df64706a0be0, points: 0, grade: false, solution: true -->
```python
def myPoly(coeffs):
    """coeffs is a list of floats or ints, function returns a polynomial with 
    the kth entry in coeffs as the coefficient of the monomial x^k"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6756057eac2876cf, points: 3, grade: true, solution: false -->
```python
assert abs(myPoly([1,2,1])(5) - 36) <= .1
assert abs(myPoly([1,1,1])(5) - 31) <= .1
assert abs(myPoly([1,2,3])(4) - 57) <= .1
```

### Exercise 6 (5 pts, manually graded)

#### Part A (2 pts, manually graded)
Let's make a new class called `House` that where you have to provide the following attributes in the `__init__` function:
- address,
- style, (see [this list](https://www.diynetwork.com/how-to/rooms-and-spaces/exterior/26-popular-architectural-home-styles-pictures) for inspiration)
- yearBuilt
- numUnits (are there rental units in the property?)
- numFloors (how many floors does it have?)

You should also include the attributes and provide methods to change the values of (or just check!) 
- doorsLocked
- windowsClosed

<!-- grade_id: cell-1cd99f2f668b3de6, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### Part B (3pts, manually graded)

Add at least 3 new attributes and 3 methods for manipulating these attributes. Get creative!

<!-- grade_id: cell-e7da709b5aeb47be, points: 3, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```