# PS06 - Donald Trump's Twitter Data (15 pts)

## Due October 31st, 2021

![Donald Trump Twitter Data Set](trump_twitter.jpeg)

### Topics Covered

This problem set is inspired by [Chris Tralie's Famous IDS 301 Python Course at Ursinus College](https://ursinus-ids301-s2020.github.io/CoursePage/Schedule/index.html). **Please do not post solutions on the internet, outside of our UAlbany ITS servers.**

In this problem set we practice "data wrangling" with a real data set that is a JSON file of Donald Trump's (old) twitter posts. These are loaded into a dictionary using Pickle. 

We then do some exploratory data analysis on this data set using our skills with NumPy and Python dictionaries.

## Question 1 (4 pts)

Load the data file into a python list using pickle. Each element in the list is a dictionary. 

### Part A (2 pts)
Complete the function below that returns the number of tweets in the data set. 

<!-- grade_id: cell-ef69b2a8c4a5d23a, points: 0, grade: false, solution: true -->
```python
import pickle

tweets = pickle.load(open("trumpSinceElection.dat", "rb"))

def how_many_tweets(tweets):
    # YOUR CODE HERE
    raise NotImplementedError()

#how_many_tweets(tweets)
```

<!-- grade_id: cell-264d66f8579ef60d, points: 2, grade: true, solution: false -->
```python
assert how_many_tweets(tweets) > 100
assert how_many_tweets(tweets) < 100000
```

### Part B (2 pts)

Complete the function below that returns the text of the first tweet.

<!-- grade_id: cell-03f20d8e7fddf805, points: 0, grade: false, solution: true -->
```python
def first_tweet(tweets):
    # YOUR CODE HERE
    raise NotImplementedError()

#first_tweet(tweets)
```

<!-- grade_id: cell-37c5293eb05ca6e4, points: 2, grade: true, solution: false -->
```python
assert first_tweet(tweets) != 'Four score and seven years ago our fathers brought forth, upon this continent, a new nation, conceived in liberty, and dedicated to the proposition that all men are created equal.'
```

## Question 2 (4 pts)

In this question we're going to create two numpy arrays from the `tweets` list, one that contains the **retweet count**, and one that contains the **favorite count**, i.e. the number of "likes".  Then we're going to use NumPy's `argmax` function to give you the index of the most retweeted tweet and the most favorited tweet.

We will then use this to complete the two functions below.

### Part A (2pts)

Write a function that returns Donald Trump's most retweeted tweet.

<!-- grade_id: cell-7394b51eca519f0d, points: 0, grade: false, solution: true -->
```python
import numpy as np

def most_retweeted(tweets):
    retweet_list = np.zeros(len(tweets))
    # YOUR CODE HERE
    raise NotImplementedError()
#most_retweeted(tweets)
```

<!-- grade_id: cell-e9f4c42d3c060a9b, points: 2, grade: true, solution: false -->
```python
assert most_retweeted(tweets) != "I love the New York Times."
```

### Part B (2pts)

Write a function that returns Donald Trump's most favorited tweet.

<!-- grade_id: cell-f70fc956af124c92, points: 0, grade: false, solution: true -->
```python
def most_favorited(tweets):
    favorite_list = np.zeros(len(tweets))
    # YOUR CODE HERE
    raise NotImplementedError()
most_favorited(tweets)
```

<!-- grade_id: cell-a956b6945304c519, points: 2, grade: true, solution: false -->
```python
assert most_favorited(tweets) != "I know that I am intelligent, because I know that I know nothing."
assert most_favorited(tweets) != "Wisdom begins in wonder."
assert most_favorited(tweets) != "Understanding a question is half an answer."
```

## Question 3 (4 pts)

Finish the function `find_kth_popular_tweet` that takes in the dictionary containing the tweets and an integer $k$ and returns the *text* of the $k$-th most "popular" tweet. "Popular" having two different meanings in each of the following parts.

**Hint:** You might want to use NumPy's `argsort` function.

Notice that the code below uses the `enumerate` function, which we have not gone over in class. You can learn about it [here.](https://www.geeksforgeeks.org/enumerate-in-python/)

### Part A (2 pts)

Assume here that "popular" means most retweeted tweet.

<!-- grade_id: cell-aecbcdb7fdeec21b, points: 0, grade: false, solution: true -->
```python
import numpy as np

def find_kth_retweeted(tweets, k):
    
    retweet_list = np.zeros(len(tweets))
    
    for index, tweet in enumerate(tweets):
        retweet_list[index] = tweet['retweet_count']
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-558505b8b816a635, points: 2, grade: true, solution: false -->
```python
assert find_kth_retweeted(tweets, 2) == 'TODAY WE MAKE AMERICA GREAT AGAIN!'
assert find_kth_retweeted(tweets, 3) == 'Why would Kim Jong-un insult me by calling me "old," when I would NEVER call him "short and fat?" Oh well, I try so hard to be his friend - and maybe someday that will happen!'
assert find_kth_retweeted(tweets, 5) == 'Such a beautiful and important evening! The forgotten man and woman will never be forgotten again. We will all come together as never before'
```

### Part B (2 pts)

Assume here that "popular" means most favorited tweet.

<!-- grade_id: cell-54232c052dd26afa, points: 0, grade: false, solution: true -->
```python
import numpy as np

def find_kth_favorited(tweets, k):
    
    favorite_list = np.zeros(len(tweets))
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0c8f068049ab80a4, points: 2, grade: true, solution: false -->
```python
assert find_kth_favorited(tweets,4) == 'MERRY CHRISTMAS!'
assert find_kth_favorited(tweets,10) == 'THE UNITED STATES LOVES ITALY! https://t.co/DESUJ4QM41'
assert find_kth_favorited(tweets,13) == 'HAPPY NEW YEAR!'
```

## Question 4 (3 pts, autograded)

Below I've written a function `get_words` that extracts out a dictionary of words coming from Donald Trump's twitter data. 



### Part A

Write a second function `get_top_k_words` that takes in the dictionary containing the tweets and an integer $k$ and returns the $k$ most used words, as a listed sorted in descending order.

<!-- grade_id: cell-b7dfc4b4d4f60add, points: 0, grade: false, solution: true -->
```python
def get_words(tweets):
    words = {}

    for index, tweet in enumerate(tweets):
        for word in tweet['text'].split(' '):
            if word in words:
                words[word] += 1
            else:
                words[word] = 1

    return words

def get_top_k_words(tweets,k):
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-36d27244d3e96f68, points: 2, grade: true, solution: false -->
```python
assert get_top_k_words(tweets,10) == ['the', 'to', 'and', 'of', 'a', 'is', 'in', 'RT', 'for', 'on']
```

### Part B

**What is the $46th$ most commonly used word by Donald Trump, according to the above data set?**

<!-- grade_id: cell-115886748e228a57, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE