# Test Problem Set (1 point)

For this assignment we test the functionality of the JupyterHub+NBgrader workflow.

**Question 1 (1 pt):**

Write a function `hello_world()` that returns the string `Hello, World!`

<!-- grade_id: cell-209501a82d04627a, points: 0, grade: false, solution: true -->
```python
def hello_world():
    ### BEGIN SOLUTION
    return "Hello, World!"
    ### END SOLUTION
```

Now we're going to have some test cases to make sure you're function does the right thing.

<!-- grade_id: cell-bd7e2df80976fd61, points: 1, grade: true, solution: false -->
```python
"""Check that sum returns the correct output"""
assert type(hello_world()) == str #This checks that the type of the returned value is a string.
assert hello_world() != "hello world" #Did you return a string exactly as it is written?
assert hello_world() != "hello, world" #Another test. The true answer is hidden!
### BEGIN HIDDEN TESTS
assert hello_world() == "Hello, World!"
### END HIDDEN TESTS
```

```python

```