# Problem Set 05 (16 pts)
## Due Sunday, October 24th, 2021
### Topics:
* Pandas
* Covariance and Correlation
* Regression

### Background

In the first iteration of AMAT 502, I polled students to answer the following questionnaire and recorded their responses in two CSV (comma separated value) files.

Here are the questions:
![Course Evaluation](course-eval.png)

**Exercise 1 (2pts, autograded, but will generate errors. DON'T WORRY ABOUT ERRORS HERE):** I've included with this assignment two `.csv` files: `Section 1.csv` and `Section 2.csv`. 

Complete the functions below to use Pandas to read in each csv file and create a DataFrame for each, called `sect1` and `sect2`, respectively. Then, use the command `.drop(columns =[  ])` to get rid of the `'Question Number'` column.

Each of the following two functions must return `sect1` and `sect2` respectively, after you've dropped the `'Question Number'` column. 

**The tests below will generate errors. This is okay, because when you upload it, the missing code necessary for compiling the code correctly will be inserted automatically.**

<!-- grade_id: cell-03984864268cb86f, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd

def q1_sect1():
    # YOUR CODE HERE
    raise NotImplementedError()
    return sect1
def q1_sect2():
    # YOUR CODE HERE
    raise NotImplementedError()
    return sect2
```

<!-- grade_id: cell-2c2744daa28b0751, points: 2, grade: true, solution: false -->
```python
from pandas.testing import assert_frame_equal
assert_frame_equal(q1_sect1(), ans1)
assert_frame_equal(q1_sect2(), ans2)
```

**Exercise 2 (2pts):** Use the .concat([ , ]) function to create a new DataFrame `sect` that has the responses to each question from every student, i.e. both section's responses have now been combined into a single dataframe.

<!-- grade_id: cell-45e4213d1a090abb, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Exercise 3 (4 pts in total)** 

**Part A (2pts):** Compute the mean and standard deviation of the first and second sections using methods built into Pandas.

<!-- grade_id: cell-ed3276913ade3671, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Part B (2pts):**  Using the interperetations of mean and standard deviation from the lecture, what can you say about how the class feels on average with this information (recall that the scores are out of 5 points)?

<!-- grade_id: cell-1a05f42c479c2696, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 4 (4pts):** Compute the covariance and correlation matrices of the 10 questions for all of the reviews (i.e. the concatenated DataFrames).

**Part A (1 pt): Compute Covariance of the `sect` dataFrame**

<!-- grade_id: cell-2fb4074e49a66c36, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Part B (1 pt): Compute Correlation of the `sect` dataFrame**

<!-- grade_id: cell-c09624087b951796, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Part C (2pts):** What two questions seem to be the most related?

<!-- grade_id: cell-053d701325ed9752, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**Exercise 5 (4pts):** Using the code given in the lecture, create a linear regression model that relates the two most related questions from the midterm review and plot the points along with the regression line. Also, print the equation of the line that best fits the data.

<!-- grade_id: cell-cfdde480c6d80ab6, points: 4, grade: true, solution: true -->
```python
import matplotlib.pyplot as plt
%matplotlib inline
import seaborn as sns; sns.set()
import numpy as np
from sklearn.linear_model import LinearRegression
model = LinearRegression(fit_intercept = True)
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```