# Problem Set 01 (14 Points)

## Due Friday, September 3rd, 2021 by midnight.

### Each question is worth 2 points

<p> 
    <u>Contents</u>
<ul>
    
<li>Variable assignment</li>

<li>Print</li>

<li>Data Types</li>

<li>Loops</li>

</ul>
<p>

**Exercise 1:** Write a function `sum(x,y)` that returns a string indicating what the sum of $x$ and $y$ is. For example, when x=3 and y=17 your function should return the string `'The sum of 3 and 17 is 20.'` 

<!-- grade_id: cell-ec0eeb3d2f39c461, points: 0, grade: false, solution: true -->
```python
def sum(x,y):
    # YOUR CODE HERE
    raise NotImplementedError()
```

Check that when x=3 and y=17, that your function returns the string 'The sum of 3 and 17 is 20.'

```python
sum(3,17)
```

<!-- grade_id: cell-f569382c691d7102, points: 2, grade: true, solution: false -->
```python
"""Check that sum returns the correct output"""
assert sum(102,-5) == "The sum of 102 and -5 is 97."
assert sum(25,10) == "The sum of 25 and 10 is 35."
```

**Exercise 2:** Finish the function below that takes in the length of the two parallel sides of a trapezoid $a$ and $b$, which are separated by a distance $h$, and returns its area.

<!-- grade_id: cell-5b94e3d5b67a9d28, points: 0, grade: false, solution: true -->
```python
def area_of_trapezoid(a, b, h):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e714d87d023ae8da, points: 2, grade: true, solution: false -->
```python
"""Check that area_of_triangle returns the correct output for several inputs"""
assert abs(area_of_trapezoid(3,3, 5) - 15) <=.1
assert abs(area_of_trapezoid(3,4, 5) - 17.5) <=.1
```

**Exercise 3:** Write a program to convert temperatures from celsius to fahrenheit where 
    $$ c = \frac{5}{9}(f-32)$$

<!-- grade_id: cell-5ca9f6058202dbd4, points: 0, grade: false, solution: true -->
```python
def c_to_f(c):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4ae7d8989344a234, points: 2, grade: true, solution: false -->
```python
"""Check that formula is correct for several inputs"""
assert c_to_f(22) == 71.6
assert c_to_f(0) == 32.0
```

**Exercise 4:** Write a function that takes an integer as input and outputs whether or not the number is odd or even. Make sure you have an output that prints "The number x is even/odd", where $x$ was the input.

<!-- grade_id: cell-8bdf535b86f81065, points: 0, grade: false, solution: true -->
```python
def even_or_odd(x):
    """Returns a string that declares if the entered number is even or odd.
    You may assume the input is an integer."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
even_or_odd(17)
```

<!-- grade_id: cell-90a84bd2ec87ba66, points: 2, grade: true, solution: false -->
```python
"""Check that formula is correct for several inputs"""
assert even_or_odd(9) == 'The number 9 is odd.'
assert even_or_odd(12) == 'The number 12 is even.'

```

**Exercise 5 (2 points - Manually Graded):** 
Write a Python function `draw(n)` that constructs the following pattern using the `print` function and a nested for loop. Here is what the output should look like for $n=5$, but with less space between lines.

\*

\* \*

\* \* \* 

\* \* \* \* 

\* \* \* \* \*

\* \* \* \* 

\* \* \* 

\* \* 

\*
<p>

<!-- grade_id: cell-de10cb1197843042, points: 2, grade: true, solution: true -->
```python
def draw(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

**Exercise 6:** Write a method that takes in two numbers and finds the quotient and remainder of the first input value, divided by the second input value. Return the string "The quotient is _. The remainder is _"  where the blanks include your answer. 


<!-- grade_id: cell-bb097361dbaa2edc, points: 0, grade: false, solution: true -->
```python
def quotient_remainder(a,b):
    """Return a string that says 'The quotient is q and the remainder is r' 
    for q and r depending on a and b.
    You may assume a and b are integers with a > b"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-6723680b81207a08, points: 2, grade: true, solution: false -->
```python
assert quotient_remainder(7,3) == 'The quotient is 2 and the remainder is 1'
assert quotient_remainder(10,3) == 'The quotient is 3 and the remainder is 1'
```

**Exercise 7:** Write a function that takes in an integer $n$ and computes $n!$. Do this without recursion.


<!-- grade_id: cell-fa29a00c195e5c2e, points: 0, grade: false, solution: true -->
```python
def factorial_iter(n):
    """Takes in an integer n>0 and returns the product of all integers from 1 to n."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python

```

<!-- grade_id: cell-21dfe4353655d5e5, points: 2, grade: true, solution: false -->
```python
assert factorial_iter(6) == 720
assert factorial_iter(7) == 5040
assert factorial_iter(10) == 3628800
```

```python

```