# Problem Set 02 (20 points possible)
### Due Sunday September 19th at 11:59 ET on JupyterHub
## Content

- **Conditionals**
- **Loops**
- **Strings**
- **Lists**

#### Nota Bene: You CANNOT use functions from the math or scipy packages that use pre-computed values of $e$ or factorial operations.


**Exercise 1 (3 pts):** Write a function that takes in an integer $n$ and prints the $n$-th Taylor approximation to the number $e$. By this we mean, complete the function below and use the code provided to sum the first $n$ terms of the series
$$e \approx 1+ 1 + \frac{1^2}{2!} + \frac{1^3}{3!} + \frac{1^4}{4!}+ \cdots $$ and so on.

<!-- grade_id: cell-4377672c617bf9ce, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    total=1
    j=1
    while j <= n:
        total=total*j
        j=j+1
    return total

def partial_e_series(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2a5a27f54a2f0cab, points: 3, grade: true, solution: false -->
```python
assert abs(partial_e_series(3)-2.666) <= .01
assert abs(partial_e_series(4)- 2.7083) <= .001
```

**Exercise 2 (3 pts):** Write a for loop that uses the integer $n$ to compute the $n^{th}$ left hand Riemann sum of the function $f(x) = x^2$ on the interval [0,1]. 

*Hint* As $n\to \infty$, what should the returned sum should tend to?

<!-- grade_id: cell-12e8cac17e45d049, points: 0, grade: false, solution: true -->
```python
def riemann_xsquared(n):
    sum=0
    # YOUR CODE HERE
    raise NotImplementedError()
    return sum
riemann_xsquared(7)
```

<!-- grade_id: cell-f7eaa60b22256a9d, points: 3, grade: true, solution: false -->
```python
assert abs(riemann_xsquared(1) - 0) <=.01
assert abs(riemann_xsquared(2) - 0.125) <= .1
assert abs(riemann_xsquared(3) - 0.185185) <= .1
assert abs(riemann_xsquared(10) - 0.285) <= .01
```

**Exercise 3 (2 pts):** Complete the function below that takes in three integers $(a,b,c)$ and returns the number of integers between $a$ and $b$ that are divisible by $c$. Make sure it tests $a$ and $b$ for divisibility by $c$.


<!-- grade_id: cell-e6426ff88577549f, points: 0, grade: false, solution: true -->
```python
def how_many_divisible(a,b,c):
    counter=0
    # YOUR CODE HERE
    raise NotImplementedError()
    return counter
```

<!-- grade_id: cell-e9dffd315102441d, points: 2, grade: true, solution: false -->
```python
assert how_many_divisible(10,20,5) == 3
assert how_many_divisible(30,77,7) == 7
```

**Exercise 4 (2 points):** Write a python program that takes in three integers and returns the largest *odd* integer.

If there are no odd integers return the string 'Sorry. All arguments are even.

<!-- grade_id: cell-c2012b30787a5202, points: 0, grade: false, solution: true -->
```python
def largest_odd(x,y,z):
    """Arguments x,y,z are assumed to be of type int. Output is largest odd integer if there is one.
    If none of the arguments are odd, return the message 'all arguments are even.'"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-42ac30dda8919a27, points: 2, grade: true, solution: false -->
```python
assert largest_odd(2,12,11) == 11
assert largest_odd(2,12,14) == 'Sorry. All arguments are even.'
assert largest_odd(-13,17,-39) == 17
```

**Exercise 5 (3 pts):** Write a program `ffd(s)` that takes in a string `s` of decimal numbers, separated by commas, e.g. `s='1.5,6.4,3.8'`, and returns the sum of the *decimal* numbers. In our example, `ffd(s)` should return `11.7`.

You may want to use the fact that `s[i,j]` returns the portion of the string starting at the character in position $i$ and ending at the character in position $j-1$. Recursion might also help you in your solution, but all that matters is that the tests below are satisfied.

<!-- grade_id: cell-350d157902442095, points: 0, grade: false, solution: true -->
```python
def ffd(s):
    """Assumes s is a string of ints or floats separated by commas and returns the sum of those ints/floats as a float."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fce6a11e6f21d21e, points: 3, grade: true, solution: false -->
```python
assert ffd('1.5,6.4,3.8') == 11.7
assert ffd('22.4,43.2') == 65.6
assert ffd('2,3,4') == 9.0
```

**Exercise 6 (3 pts):** Write a function that takes in a list of lists of equal length and then produces their *average* list, i.e. this is a list where the entry in index `i` is the average of all the given entries in index `i`.

<!-- grade_id: cell-e56fa423dbe126d9, points: 0, grade: false, solution: true -->
```python
def avgList(myList):
    """Assumes the argument myList is a list of lists, 
    all of the same length k with int or float entries. 
    Returns a list of length k that computes the entrywise average."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e91c52c007d1f919, points: 3, grade: true, solution: false -->
```python
test1 = [2.0, 4.5, 6.5, 8.0]
for j in [abs(test1[i] - avgList([[3,5,7,8],[1,4,6,8]])[i]) <.1 for i in range(4)]: 
    assert j
test2 = [2.0, 2.0, 2.0]
for j in [abs(test2[i] - avgList([[1, 2, 3], [1, 3, 2], [2, 1, 3], [2, 3, 1], [3, 1, 2], [3, 2, 1]])[i]) <.1 for i in range(3)]: 
    assert j
```

**Exercise 7: (4 pts total, manually graded)**

**Part (a):** Write a function that takes in a list of n numbers and returns a list of all possible permutations of that list.

**Part (b):** Compose the function you wrote in Part (a) with the function you wrote in Exercise 6 to compute the average of the list `[1,3,5,7]` under all 24 of its permutations.

<!-- grade_id: cell-2546ecd8a50f49d9, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```