# Problem Set 02
## Quiz 1 is on September 11th, 2024

- **Conditionals**
- **Loops**
- **Lists**
- **Recursion**

#### Nota Bene: For this PS and Quiz 1, functions from the math or scipy packages that use pre-computed values of $e$ or factorial operations are PROHIBITED


**Exercise 1 (2 pts, autograded):** Complete the function below that takes in three integers $(a,b,c)$ and returns the number of integers between $a$ and $b$ (inclusive of the endpoints) that are divisible by $c$. Make sure it tests $a$ and $b$ for divisibility by $c$!


<!-- grade_id: cell-e6426ff88577549f, points: 0, grade: false, solution: true -->
```python
def how_many_divisible(a,b,c):
    counter=0
    # YOUR CODE HERE
    raise NotImplementedError()
    return counter
```

<!-- grade_id: cell-e9dffd315102441d, points: 2, grade: true, solution: false -->
```python
assert how_many_divisible(10,20,5) == 3
assert how_many_divisible(30,77,7) == 7
```

**Exercise 2 (3 pts, autograded):** Write a function that takes in an integer $n$ and returns the $n$-th Taylor approximation to the number $e$. By this we mean, complete the function below and use the code provided to sum the first $n$ terms of the series
$$e \approx 1+ 1 + \frac{1^2}{2!} + \frac{1^3}{3!} + \frac{1^4}{4!}+ \cdots + \frac{1^n}{n!} + \cdots $$

Use Python based indexing, i.e. when $n=0$ it should return $1$, when $n=1$ it should return 2, when $n=2$ it should return 2.5, and so on.

<!-- grade_id: cell-4377672c617bf9ce, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    total=1
    j=1
    while j <= n:
        total=total*j
        j=j+1
    return total

def partial_e_series(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2a5a27f54a2f0cab, points: 3, grade: true, solution: false -->
```python
assert abs(partial_e_series(3)-2.666) <= .01
assert abs(partial_e_series(4)- 2.7083) <= .001
```

**Exercise 3 (3 pts, autograded):** Write a for loop that uses the integer $n$ to compute the $n^{th}$ left hand Riemann sum of the function $f(x) = x^2$ on the interval [0,1]. 

*Hint* As $n\to \infty$, what should the returned sum should tend to?

<!-- grade_id: cell-12e8cac17e45d049, points: 0, grade: false, solution: true -->
```python
def riemann_xsquared(n):
    sum=0
    # YOUR CODE HERE
    raise NotImplementedError()
    return sum
#riemann_xsquared(7)
```

```python
riemann_xsquared(3)
```

```python
(1/9)*(1/3)+(4/9)*(1/3)
```

<!-- grade_id: cell-f7eaa60b22256a9d, points: 3, grade: true, solution: false -->
```python
assert abs(riemann_xsquared(1) - 0) <=.01
assert abs(riemann_xsquared(2) - 0.125) <= .1
assert abs(riemann_xsquared(3) - 0.185185) <= .1
assert abs(riemann_xsquared(10) - 0.285) <= .01
```

**Exercise 4 (6 points, manually graded):**

**Newton-Raphson Method from Lec03-LIVE**

***Part A (4 of the 6 pts)***

Write three functions.
1. A function `p(x)` that returns the value $x^3 -10x^2 + 5x + 2$
2. A function `dp(x)` that returns the value $3x^2 - 20x + 5$
3. A function `rootsOfp(startguess)` that returns the solution of $x^3 -10x^2 + 5x + 2 =0$ within .001 accuracy, by using Newton's method, i.e. you start with `guess=startguess` and then update `guess = guess - p(guess)/dp(guess)` while `p(guess)>.001`.


<!-- grade_id: cell-92ceb5ce8183eec3, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

***Exercise 4, Part B (2 of the 6 pts)***

How many roots are there? What are they? Give examples of values of `startguess` that converge to those roots.

<!-- grade_id: cell-9077ace6cb07ab9a, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**Exercise 5 (2 pts):** Write a function that takes in a list of lists of equal length and then produces their *average* list, i.e. this is a list where the entry in index `i` is the average of all the given entries in index `i`.

<!-- grade_id: cell-e56fa423dbe126d9, points: 0, grade: false, solution: true -->
```python
def avgList(myList):
    """Assumes the argument myList is a list of lists, 
    all of the same length k with int or float entries. 
    Returns a list of length k that computes the entrywise average."""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e91c52c007d1f919, points: 2, grade: true, solution: false -->
```python
test1 = [2.0, 4.5, 6.5, 8.0]
for j in [abs(test1[i] - avgList([[3,5,7,8],[1,4,6,8]])[i]) <.1 for i in range(4)]: 
    assert j
test2 = [2.0, 2.0, 2.0]
for j in [abs(test2[i] - avgList([[1, 2, 3], [1, 3, 2], [2, 1, 3], [2, 3, 1], [3, 1, 2], [3, 2, 1]])[i]) <.1 for i in range(3)]: 
    assert j
```

**Exercise 6: (4 pts total, manually graded)**

**Part (a):** Write a function that takes in a list of n numbers and returns a list of all possible permutations of that list. **DO NOT USE ITERTOOLS**

**Part (b):** Compose the function you wrote in Part (a) with the function you wrote in Exercise 6 to compute the average of the list `[1,3,5,7]` under all 24 of its permutations.

<!-- grade_id: cell-2546ecd8a50f49d9, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```